## 6. Funciones

Topicos:
1. Como declarar una función
2. Duck-typing en Julia
3. Mutating vs. non-mutating functions
4. Algunas funciones de mayor órden

### Como declarar una función

Julia tiene varias formas de escribir una función. La primera requiere de los comandos `function` y `end`

```julia
function saluda(nombre)
    println("Hola $nombre, que lindo verte!")
end
```

```julia
function f(x)
    x^2
end
```

Podemos aplicar cualquiera de estas dos funciones así:

```julia
saluda("Juana")
```

```julia
f(42)
```

Alternativamente, podríamos haber definido las funciones en una única línea:

```julia
saluda2(nombre) = println("Hola $nombre, que lindo verte!")
```

```julia
f2(x) = x^2
```

```julia
saluda2("Juan")
```

```julia
f2(42)
```

Por último, podríamos haberlas declarado como funciones "anónimas" (más abajo veremos para que sirven).

```julia
saluda3 = name -> println("Hola $name, que lindo verte!")
```

```julia
f3 = x -> x^2
```

```julia
saluda3("José")
```

```julia
f3(42)
```

### Duck-typing en Julia

*"If it walks like a duck and it quacks like a duck, it's a duck."*

Las funciones de Julia funcionan sobre cualquier input que tenga sentido.

`saluda` puede saludar a un número:

```julia
saluda(55595472)
```

Y `f` funciona con matrices:

```julia
A = rand(3, 3)
A
```

```julia
f(A)
```

`f` también funciona con una string como "hola" porque `*` está definido para strings como concatenación:

```julia
f("hola")
```

Por otro lado, `f` no funciona sobre vectores. A diferencia de `A^2`, que está bien definido, el significado de `v^2` para un vector `v` no es una operación algebráica bien definida.

```julia
v = rand(3)
```

```julia
# This won't work
f(v)
```

## Mutating vs. non-mutating functions

Por convención, las funciones que terminan con `!` pueden modificar el contenido de sus inputs, y las funciones que no tienen `!` no lo hacen.

Por ejemplo, veamos la diferencia entre `sort` y `sort!`.


```julia
v = [3, 5, 2]
```

```julia
sort(v)
```

```julia
v
```

`sort(v)` devuelve un arreglo ordenado que contiene los mismos elementos que `v`, pero `v` no se modifica.

Por otro lado, cuando corremos `sort!(v)`, los contenidos de `v` son ordenados **modificando** `v`.

```julia
sort!(v)
```

```julia
v
```

### Algunas funciones de mayor orden

#### `map`

`map` es una función de "mayor orden" en Julia que toma una *función* como uno de sus inputs.

`map` luego aplica esa función a cada elemento de la estructura de datos que le pasamos como segundo argumento. Por ejemplo, correr:

```julia
map(f, [1, 2, 3])
```

va a resultar en un arreglo donde la función `f` se aplicó a todos los elementos de `[1, 2, 3]`

```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Elevamos al cuadrado todos los elementos del vector `[1, 2, 3]`, en lugar de elevar el vector `[1, 2, 3]` al cuadrado.

Podríamos haber hecho lo mismo con una función anónima, en lugar de una función con nombre como `f`. 

Por ejemplo, podemos elevar los elementos de `[1, 2, 3]` al cubo de la siguiente forma:

```julia
x -> x^3
```

via

```julia
map(x -> x^3, [1, 2, 3])
```

and now we've cubed all the elements of `[1, 2, 3]`!

#### `broadcast`

`broadcast` es una generalización de `map`. Puede hacer lo mismo que `map` y más (ahora veremos qué). La sintaxis es igual a `map`

```julia
broadcast(f, [1, 2, 3])
```

nuevamente, aplicamos `f` a todos los elementos de `[1, 2, 3]` - esta vez *broadcasting* `f`.

Otra forma de llamar la función `broadcast` es poniendo un `.` entre el nombre de la función que queremos `broadcast`ear  y sus inputs. Por ejemplo,

```julia
broadcast(f, [1, 2, 3])
```

es lo mismo que

```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Notar nuevamente lo distinto que es esto a correr 

```julia
f([1, 2, 3])
```

Podemos elevar al cuadrado cada elemento de un vector, pero no podemos elevar un vector al cuadrado.

Para fijar ideas, veamos la diferencia entre:

```julia
f(A)
```
y
```julia
f.(A)
```
para una matriz `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

```julia
f(A)
```

Como antes, vemos que para una matriz, `A`,

```
f(A) = A^2 = A * A
``` 

Por otro lado,

```julia
B = f.(A)
```

contiene todos los elementos de `A` elevados al cuadrado.

Esta "dot syntax" para hacer broadcasting nos permite escribir expresiones relativamente complejas de una forma que se lee naturalmente y es más cercana a la notación matemática. Por ejemplo, podemos escribir:

```julia
A .+ 2 .* f.(A) ./ A
```

en lugar de

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

y ambas funcionarán exactamente igual. 

### Ejercicios

#### 6.1 

Escribir una función `suma_uno` que agregue 1 a su input.

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
@assert add_one(11) == 12
```

#### 6.2 

Usar `map` o `broadcast` para incrementar cada elemento de `A` en `1` y asignar el resultado a la matriz `A1`.

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 6.3 

Usar la sintaxis de `broadcast` con `.` para incrementar cada elemento de la matriz `A1` en 1 y guardarlo en la matriz `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```